# Exercise: Decorators 1
# I AM NOT DONE
#
# `call_it` is a higher-order function: it takes a function as an argument
# and calls it, returning whatever the function returns.
#
# Fill in the body — replace `pass` with the correct return statement.


def call_it(func):
    pass


def say_hello():
    return "hello"


def give_seven():
    return 7
