# Exercise: Variables 1
# ----------------------
# I AM NOT DONE
#
# Replace each `???` with a value of the right type so the checks pass.
#   a should be an int with value 0
#   b should be a float with value 0.0
#   c should be the empty string

a = ???
b = ???
c = ???

# --- checks (do not edit below) ---
assert isinstance(a, int), "a should be an integer"
assert isinstance(b, float), "b should be a float"
assert isinstance(c, str), "c should be a string"
assert a == 0
assert b == 0.0
assert c == ""
print("variables1 ✓")
