# Exercise: Variables 5
# I AM NOT DONE
#
# Use augmented assignment operators to update `total`:
#   Start with total = 10, then add 5, then subtract 3, then multiply by 4.
# Fill in the ??? placeholders with the correct augmented-assignment operators.

total = 10
total ???= 5
total ???= 3
total ???= 4
