# Exercise: Variables 8
# I AM NOT DONE
#
# Use type() to inspect a value and type-conversion functions to create new values:
#   - `kind` should be the type of 3.14 (i.e. type(3.14))
#   - `as_int` should be 3.14 converted to an int using int()
#   - `as_str` should be 42 converted to a string using str()
# Replace each ??? with the correct expression.

kind = ???
as_int = ???
as_str = ???
