# Exercise: Variables 10
# I AM NOT DONE
#
# Combine everything: given width and height of a rectangle, compute:
#   area      = width * height
#   perimeter = 2 * (width + height)
#   diagonal  = (width**2 + height**2) ** 0.5   (Pythagorean theorem)
# width = 3, height = 4.  Replace each ??? with the correct expression.

width = 3
height = 4
area = ???
perimeter = ???
diagonal = ???
