# Exercise: Lists 8
# I AM NOT DONE
#
# Given `temperatures = [0, 100, -40, 37, 212, -273]` (in Celsius),
# build a new list `fahrenheit` that contains each temperature converted
# to Fahrenheit using the formula:  F = C * 9 / 5 + 32
#
# Use a loop (or a list comprehension) to build the list.
# Do not modify `temperatures`.

temperatures = [0, 100, -40, 37, 212, -273]

fahrenheit = []
for temp in temperatures:
    pass  # replace this line with the conversion and append
