# Exercise: Lists 7
# I AM NOT DONE
#
# Given `fruits = ["apple", "banana", "cherry", "banana", "apple", "apple"]`:
#   `total`   — total number of elements (use `len()`)
#   `has_kiwi` — True if "kiwi" is in the list, False otherwise (use `in`)
#   `apple_count` — how many times "apple" appears (use `.count()`)

fruits = ["apple", "banana", "cherry", "banana", "apple", "apple"]

total       = ???
has_kiwi    = ???
apple_count = ???
