# Exercise: Lists 6
# I AM NOT DONE
#
# You have `scores = [42, 7, 99, 55, 3, 76]`.
# 1. Sort `scores` in-place in ascending order using `.sort()`.
# 2. Use `sorted()` to create a new list `desc` with the same values
#    sorted in descending order (scores itself must stay ascending).

scores = [42, 7, 99, 55, 3, 76]

# Step 1: sort scores in-place (ascending)
???

# Step 2: build `desc` — descending order, without modifying scores again
desc = ???
