# Exercise: Lists 3
# I AM NOT DONE
#
# Start with `items = ["apple", "banana"]`.
# 1. Append "cherry" to the end of the list.
# 2. Insert "mango" at index 1 (so it becomes the second element).
# After both operations the list should be:
#   ["apple", "mango", "banana", "cherry"]

items = ["apple", "banana"]

# Step 1: append "cherry"
???

# Step 2: insert "mango" at index 1
???
