# Exercise: Lists 2
# I AM NOT DONE
#
# Fix the code below so that:
#   - `first` holds the first element of `numbers` (should be 10)
#   - the element at index 2 is changed from 30 to 99
#
# Only fix the two lines marked with ???.

numbers = [10, 20, 30, 40, 50]

first = numbers[???]
numbers[???] = 99
