# Exercise: Lists 4
# I AM NOT DONE
#
# Given `items = ["apple", "banana", "cherry", "date"]`:
# 1. Remove "banana" from the list using `.remove()`.
# 2. Pop the last element and store it in `last`.
# After both operations `items` should be ["apple", "cherry"]
# and `last` should be "date".

items = ["apple", "banana", "cherry", "date"]

# Step 1: remove "banana"
???

# Step 2: pop the last element into `last`
last = ???
