# Exercise: Lists 5
# I AM NOT DONE
#
# Given `letters = ["a", "b", "c", "d", "e", "f"]`, use slicing to
# create three new lists:
#   `first_three` — the first three elements: ["a", "b", "c"]
#   `last_two`    — the last two elements:    ["e", "f"]
#   `middle`      — elements at index 2 through 4 (inclusive): ["c", "d", "e"]
#
# Use slice notation (no loops, no indexing one element at a time).

letters = ["a", "b", "c", "d", "e", "f"]

first_three = letters[???]
last_two    = letters[???]
middle      = letters[???]
