# Exercise: Conditionals 10
# I AM NOT DONE
#
# Putting it all together: classify a number using several conditions.
#
# Complete `classify(n)` so it returns:
#   "large positive"   when n > 100
#   "small positive"   when n > 0  (and not > 100)
#   "zero"             when n == 0
#   "small negative"   when n >= -100  (and not 0 or positive)
#   "large negative"   when n < -100
#
# Use an if / elif / elif / elif / else chain.

def classify(n):
    if ???:
        return "large positive"
    ???
