# Exercise: Conditionals 4
# I AM NOT DONE
#
# Python comparison operators: <, <=, >, >=, ==, !=
#   ==  equal to          !=  not equal to
#   <   less than         <=  less than or equal
#   >   greater than      >=  greater than or equal
#
# Complete `compare(a, b)` so it returns:
#   "less"    when a < b
#   "equal"   when a == b
#   "greater" when a > b

def compare(a, b):
    if a ??? b:
        return "less"
    elif a ??? b:
        return "equal"
    else:
        return "greater"
