# Exercise: Conditionals 8
# I AM NOT DONE
#
# Python has "truthiness": values can be used directly in conditions
# without an explicit == check.
#
# Falsy values (treated as False in a condition):
#   0, 0.0, "", [], {}, None
# Everything else is truthy.
#
# Complete `describe(value)` so it returns:
#   "empty" when value is falsy
#   "has content" when value is truthy
#
# Do NOT use == in your condition — use `if value` directly.

def describe(value):
    if ???:
        return "has content"
    ???
