# Exercise: Conditionals 1
# I AM NOT DONE
#
# Python's simplest conditional: if the condition is True, the indented
# block runs; otherwise nothing happens.
#
# `is_sunny` is True. Write an `if` statement so that `message` is set
# to "Let's go outside!" when `is_sunny` is True.
# Leave `message` as the empty string "" before the if-block.

is_sunny = True
message = ""

if ???:
    message = "Let's go outside!"
