# Exercise: Conditionals 6
# I AM NOT DONE
#
# `not` flips a boolean: `not True` is False, `not False` is True.
# It is useful when you want to act on the absence of something.
#
# Complete `is_closed(is_open)` so it returns True when the shop is NOT open.
# Complete `neither_zero(a, b)` so it returns True when NEITHER a nor b is 0.
# Hint: use `not` in both functions.

def is_closed(is_open):
    return ???

def neither_zero(a, b):
    return ???
