# Exercise: Conditionals 2
# I AM NOT DONE
#
# An if/else gives two branches: one when the condition is True,
# another when it is False.
#
# Complete the function `sign(n)` so it returns "positive" when n > 0
# and "non-positive" otherwise.

def sign(n):
    if n > 0:
        return "positive"
    ???
