# Exercise: Dictionaries 7
# I AM NOT DONE
#
# Iterate over a dictionary to compute a summary.
#
# Given `prices` below, build `discounted` — a new dict where every
# value is reduced by 10% (multiply each price by 0.9).
# Round each result to 2 decimal places.
#
# Expected:
#   discounted == {"apple": 0.9, "milk": 1.35, "bread": 2.25}

prices = {"apple": 1.0, "milk": 1.5, "bread": 2.5}

discounted = {}
for ???:
    ???
