# Exercise: Dictionaries 4
# I AM NOT DONE
#
# Remove keys from a dict using two approaches:
#   1. Use `del` to remove "banana" from `basket`.
#   2. Use `.pop()` to remove "cherry" from `basket` and store the
#      removed value in `cherry_count`.

basket = {"apple": 5, "banana": 3, "cherry": 12, "date": 7}

# Remove "banana" with del
???

# Remove "cherry" with .pop(), storing its value in cherry_count
cherry_count = ???
