# Exercise: Dictionaries 10
# I AM NOT DONE
#
# Build a frequency dict — count how many times each word appears in `words`.
#
# Store the result in `freq`.  Keys are words, values are their counts.
#
# Expected:
#   freq == {"the": 3, "cat": 2, "sat": 1, "on": 1, "mat": 1, "and": 1, "hat": 1}
#
# Hint: start with an empty dict and use .get() or setdefault() inside a loop.

words = ["the", "cat", "sat", "on", "the", "mat", "and", "the", "cat", "hat"]

freq = {}
for ???:
    ???
