# Exercise: Dictionaries 8
# I AM NOT DONE
#
# Use the `in` operator to test key membership in a dict.
#
# Given `permissions` below, set:
#   `can_read`    — True if "read"   is a key in permissions, False otherwise
#   `can_write`   — True if "write"  is a key in permissions, False otherwise
#   `can_delete`  — True if "delete" is a key in permissions, False otherwise
#
# Do NOT hard-code True/False — use `in` expressions.

permissions = {"read": True, "write": False}

can_read = ???
can_write = ???
can_delete = ???
