# Exercise: Dictionaries 5
# I AM NOT DONE
#
# Given `inventory` below, produce three collections:
#   `all_items`  — a list of all keys (the item names), in any order
#   `all_counts` — a list of all values (the stock counts), in any order
#   `all_pairs`  — a list of all (key, value) tuples, in any order
#
# Use .keys(), .values(), and .items() respectively, converting each to a list.

inventory = {"pen": 10, "notebook": 5, "eraser": 8}

all_items = ???
all_counts = ???
all_pairs = ???
