# Exercise: Dictionaries 3
# I AM NOT DONE
#
# You have an existing dict `scores`. Do two things:
#   1. Add a new key "diana" with value 91.
#   2. Update "alice"'s value to 88 (she re-took the test).
#
# Modify `scores` in place — do not create a new dict.

scores = {"alice": 75, "bob": 60, "carol": 82}

# Add "diana" -> 91 here
???

# Update "alice" -> 88 here
???
