# Exercise: Strings 6
# I AM NOT DONE
#
# Use .strip() to remove leading and trailing whitespace from `raw` and
# assign the result to `clean`. Then use .replace() to swap every "cat"
# in `clean` with "dog" and assign the result to `replaced`.

raw = "  I love cats and cats  "
clean = ???
replaced = ???
