# Exercise: Strings 9
# I AM NOT DONE
#
# Use the `in` operator to check whether "fox" appears in `sentence` and
# assign the boolean result to `has_fox`. Then use .find() to find the
# index where "fox" starts and assign it to `fox_index`.

sentence = "the quick brown fox jumps"
has_fox = ???
fox_index = ???
