# Exercise: Functional 4
# I AM NOT DONE
#
# Use `filter()` to build `evens` — a list containing only the even numbers
# from `numbers`.
# `filter(func, iterable)` keeps elements where `func(element)` is truthy.
# Remember to wrap the result in `list(...)`.

numbers = [1, 2, 3, 4, 5, 6, 7, 8, 9, 10]

evens = list(filter(???, numbers))
