# Exercise: Functional 3
# I AM NOT DONE
#
# Use `map()` to build `doubled` — a list where every number in `numbers`
# is multiplied by 2.
# `map(func, iterable)` applies `func` to each element.
# Remember to wrap the result in `list(...)`.

numbers = [1, 2, 3, 4, 5]

doubled = list(map(???, numbers))
