# Exercise: Functional 5
# I AM NOT DONE
#
# Sort `words` by the LENGTH of each word (shortest first).
# Use `sorted()` with the `key=` parameter.
# The key function receives one word and should return its length.
# Replace ??? with a lambda that returns the length of a word.

words = ["banana", "fig", "apple", "kiwi", "date"]

by_length = sorted(words, key=???)
