# Exercise: Functional 6
# I AM NOT DONE
#
# `apply_twice` should call `func` on `value`, then call `func` on that result.
# In other words: apply_twice(f, x) → f(f(x))
# Fill in the function body — replace `pass` with the correct return statement.


def apply_twice(func, value):
    pass


triple = lambda x: x * 3
add_ten = lambda x: x + 10
