# Exercise: Functional 9
# I AM NOT DONE
#
# Use the built-in `any()` and `all()` to fill in two expressions below.
#
# `has_negative` should be True if ANY number in `numbers` is negative.
# `all_positive` should be True if ALL numbers in `positives` are positive (> 0).
#
# Replace each ??? with the correct generator expression inside any()/all().

numbers = [3, -1, 4, 1, 5]
positives = [2, 7, 11, 5]

has_negative = any(???)
all_positive = all(???)
