# Exercise: Functional 7
# I AM NOT DONE
#
# `make_multiplier(n)` should RETURN a function that multiplies its argument by n.
# Example: triple = make_multiplier(3); triple(4) → 12
#
# Replace `pass` inside `make_multiplier` with a return statement that
# returns a lambda (or inner function) capturing `n`.


def make_multiplier(n):
    pass
