# Exercise: Recursion 3
# ---------------------
# I AM NOT DONE
#
# Complete `recursive_sum(lst)` so it returns the sum of all numbers in
# `lst` using recursion.  Base case: an empty list has sum 0.  Recursive
# case: the first element plus recursive_sum of the rest.

def recursive_sum(lst):
    if not lst:
        return 0
    return ???
