# Exercise: Recursion 6
# ---------------------
# I AM NOT DONE
#
# Complete `reverse_string(s)` so it returns the string `s` reversed
# using recursion.  Base case: a string of length 0 or 1 is already
# reversed.  Recursive case: reverse the tail, then append the first
# character.

def reverse_string(s):
    if len(s) <= 1:
        return s
    return ???
