# Exercise: Recursion 1
# ---------------------
# I AM NOT DONE
#
# Complete `countdown(n)` so it returns a list counting down from n to 0
# using recursion.  Base case: when n is 0, return [0].  Recursive case:
# return [n] + countdown(n - 1).

def countdown(n):
    if n == 0:
        return [0]
    return ???
