# Exercise: Sets 7
# I AM NOT DONE
#
# Remove duplicates from a list by converting it to a set.
#
# Given `numbers` (a list with repeated values), create `unique` — a set
# containing exactly the distinct values from `numbers`.
#
# Then create `unique_count` — the number of unique values (an int).
#
# Expected: unique == {1, 2, 3, 4, 5} and unique_count == 5

numbers = [1, 2, 3, 2, 4, 1, 5, 3, 2, 4]

unique = ???
unique_count = ???
