# Exercise: Sets 4
# I AM NOT DONE
#
# Compute the union of two sets — every element that appears in either set.
#
# Given `a` and `b`, store their union in `combined`.
# Use the `|` operator or the `.union()` method.
#
# Expected: combined == {1, 2, 3, 4, 5, 6}

a = {1, 2, 3, 4}
b = {3, 4, 5, 6}

combined = ???
