# Exercise: Sets 2
# I AM NOT DONE
#
# Start with the set `colors` below.
# 1. Add the string "purple" to it using `.add()`.
# 2. Remove "orange" from it using `.discard()`.
#
# After both operations, `colors` should equal {"red", "blue", "purple"}.

colors = {"red", "blue", "orange"}
# your code here
???
