# Exercise: Sets 10
# I AM NOT DONE
#
# Compare two collections and extract meaningful subsets.
#
# You have two lists of students who attended each day of a workshop.
# Compute:
#   attended_both  — a set of students who came on BOTH days
#   attended_any   — a set of students who came on AT LEAST ONE day
#   only_day1      — a set of students who came on day 1 but NOT day 2
#   only_day2      — a set of students who came on day 2 but NOT day 1
#
# Use set operations (|, &, -) on sets converted from the lists.

day1 = ["alice", "bob", "carol", "dave"]
day2 = ["bob", "carol", "eve", "frank"]

attended_both = ???
attended_any = ???
only_day1 = ???
only_day2 = ???
