# Exercise: Sets 3
# I AM NOT DONE
#
# Use the `in` operator to check membership in a set.
#
# Given the set `fruits`, set:
#   has_apple  = True if "apple"  is in fruits, False otherwise
#   has_grape  = True if "grape"  is in fruits, False otherwise
#
# Do NOT use literals True or False directly — use the `in` operator.

fruits = {"apple", "banana", "cherry"}

has_apple = "apple" ??? fruits
has_grape = "grape" ??? fruits
