# Exercise: Sets 8
# I AM NOT DONE
#
# Test subset and superset relationships between sets.
#
# Given `small` and `big` below, compute:
#   is_subset   — True if `small` is a subset of `big`   (use <= or .issubset())
#   is_superset — True if `big` is a superset of `small` (use >= or .issuperset())
#
# Expected: both is_subset and is_superset should be True.

small = {2, 4}
big = {1, 2, 3, 4, 5}

is_subset = ???
is_superset = ???
