# Exercise: Sets 6
# I AM NOT DONE
#
# Compute the difference of two sets — elements in `a` that are NOT in `b`.
#
# Given `a` and `b`, store `a` minus `b` in `only_in_a`.
# Use the `-` operator or the `.difference()` method.
#
# Expected: only_in_a == {1, 2}

a = {1, 2, 3, 4}
b = {3, 4, 5, 6}

only_in_a = ???
