# Exercise: Sets 5
# I AM NOT DONE
#
# Compute the intersection of two sets — only elements present in BOTH sets.
#
# Given `a` and `b`, store their intersection in `shared`.
# Use the `&` operator or the `.intersection()` method.
#
# Expected: shared == {3, 4}

a = {1, 2, 3, 4}
b = {3, 4, 5, 6}

shared = ???
