# Exercise: Sets 9
# I AM NOT DONE
#
# Use a set comprehension to build a set of squares of even numbers.
#
# Given `numbers`, create `even_squares` — a set containing the square of
# every even number in `numbers`.
#
# Use a set comprehension: {expr for item in iterable if condition}
#
# Expected: even_squares == {4, 16, 36}

numbers = [1, 2, 3, 4, 5, 6, 7]

even_squares = ???
