# Exercise: Loops 2
# I AM NOT DONE
#
# Use a `for` loop over the list `fruits` to build a new list `upper_fruits`
# containing each fruit name converted to uppercase.
# Example: ["apple", "banana"] → ["APPLE", "BANANA"]

fruits = ["apple", "banana", "cherry"]
upper_fruits = []
# write your for loop here
