# Exercise: Loops 5
# I AM NOT DONE
#
# Use a `while` loop and `break` to find the first number in `values` that
# is greater than 10. Store it in `found`.
# If no number is greater than 10, leave `found` as None.
# Use an index variable `i` starting at 0, advance it each iteration,
# and `break` as soon as you find the target.

values = [3, 7, 2, 15, 8, 22]
found = None
i = 0
# write your while loop here
