# Exercise: Loops 9
# I AM NOT DONE
#
# Loop over `scores` using `.items()` to find the highest score.
# Store the name of the top scorer in `top_name` and their score in `top_score`.
# Assume all scores are non-negative integers and there is exactly one highest.

scores = {"Alice": 88, "Bob": 95, "Carol": 72, "Dave": 95 - 1}
top_name = None
top_score = -1
# write your for loop here
