# Exercise: Loops 8
# I AM NOT DONE
#
# Use `enumerate` to loop over `words` and build a list of strings where
# each entry is formatted as "N: word" (N is the 1-based index).
# Example: words = ["cat", "dog"] → ["1: cat", "2: dog"]
# Store the results in `labeled`.

words = ["python", "is", "fun"]
labeled = []
# write your for loop using enumerate here
