# Exercise: Loops 4
# I AM NOT DONE
#
# Use a `for` loop over `numbers` to compute two accumulators:
#   - `total`   : the sum of all numbers  (should be 55)
#   - `product` : the product of all numbers  (should be 3628800)
# Start `total` at 0 and `product` at 1 (the identity for multiplication).

numbers = [1, 2, 3, 4, 5, 6, 7, 8, 9, 10]
total = 0
product = 1
# write your for loop here
