# Exercise: Modules 8
# --------------------
# I AM NOT DONE
#
# Combine multiple stdlib imports to solve a small task:
#
# 1. Import `math` and `statistics`.
# 2. Use `random` (already imported below) — seed it with 0, then generate
#    a list of 10 random integers between 1 and 100 (inclusive) called `data`.
# 3. Compute `data_mean` as the mean of `data` using `statistics.mean`.
# 4. Compute `data_max_sqrt` as the square root of the maximum element of
#    `data` using `math.sqrt(max(data))`.
#
# Seeding with 0 makes results deterministic so the check can verify them.

import random

# add more imports here

random.seed(0)
data = ???
data_mean = ???
data_max_sqrt = ???
