# Exercise: Modules 3
# --------------------
# I AM NOT DONE
#
# Import the `math` module with the alias `m` using `import ... as ...`,
# then use that alias to call `m.pi` and `m.floor`.
#
# Fill in the import and complete the two expressions.

# add the import here

circle_area = ???.pi * 5 ** 2   # area of a circle with radius 5
floored = ???.floor(7.9)        # floor of 7.9
