# Exercise: Modules 5
# --------------------
# I AM NOT DONE
#
# The `string` module exposes useful constants.
# Import `string` and use its constants to create:
#
#   digits    — all digit characters ('0123456789')
#   letters   — all ASCII letters (lower + upper, i.e. `string.ascii_letters`)
#
# Assign each variable using the appropriate `string` constant directly
# (do NOT construct them manually).

# add the import here

digits = ???
letters = ???
