# Exercise: Modules 7
# --------------------
# I AM NOT DONE
#
# The `if __name__ == "__main__":` guard lets you write code that only
# runs when the file is executed directly, not when it is imported.
#
# Below, `ran_as_script` must be True only when this file is run as a
# script (i.e. when __name__ == "__main__").
#
# Replace the ??? so that `ran_as_script` is set correctly inside the guard.

ran_as_script = False

if __name__ == "__main__":
    ran_as_script = ???
