# Exercise: Modules 6
# --------------------
# I AM NOT DONE
#
# Every Python module has a built-in variable called `__name__`.
# When a file is run directly (as a script), `__name__` equals `"__main__"`.
# When a file is imported by another module, `__name__` equals the module name.
#
# Assign `module_name` to the value of `__name__` as it would appear
# when this file is run directly as a script.

module_name = ???
