# Exercise: Modules 4
# --------------------
# I AM NOT DONE
#
# Use the `random` module to pick a random integer.
#
# 1. Import the `random` module.
# 2. Seed the random number generator with `random.seed(42)`.
# 3. Assign `pick` to a random integer between 1 and 10 inclusive using
#    `random.randint(1, 10)`.
#
# Seeding makes the output deterministic so the check can verify the value.

# add the import here

# seed the generator with 42

pick = ???
