# Exercise: Comprehensions 8
# I AM NOT DONE
#
# Use a list comprehension with a conditional expression (ternary) to build
# `labels`: for each number in `values`, use "even" if the number is even,
# otherwise use "odd".
#
# Conditional expression inside a comprehension:
# [value_if_true if condition else value_if_false for item in iterable]
# Example: ["pos" if x > 0 else "non-pos" for x in [-1, 2, 0]]
#   →  ['non-pos', 'pos', 'non-pos']

values = [1, 2, 3, 4, 5, 6]

labels = ???
