# Exercise: Comprehensions 7
# I AM NOT DONE
#
# Use a list comprehension over the characters of a string to build `vowels`:
# a list of every vowel (a, e, i, o, u) found in `sentence` (lowercase).
#
# You can iterate over a string character by character:
# Example: [c for c in "hello" if c in "aeiou"]  →  ['e', 'o']

sentence = "list comprehensions are quite powerful"

vowels = ???
