# Exercise: Comprehensions 10
# I AM NOT DONE
#
# Build a lookup table `word_index` using a dict comprehension that maps each
# word in `words` to its position (index) in the list.
#
# `enumerate(iterable)` yields (index, value) pairs — use it with a dict
# comprehension to pair each word with its position:
# {word: idx for idx, word in enumerate(words)}
#
# Expected: {"apple": 0, "banana": 1, "cherry": 2, "date": 3}

words = ["apple", "banana", "cherry", "date"]

word_index = ???
