# Exercise: Comprehensions 2
# I AM NOT DONE
#
# Use a list comprehension with an `if` filter to build `evens`: only the
# even numbers from `values`.
#
# A filtered comprehension: [expr for item in iterable if condition]
# Example: [x for x in range(6) if x > 3]  →  [4, 5]

values = [1, 2, 3, 4, 5, 6, 7, 8]

evens = ???
