# Exercise: Comprehensions 4
# I AM NOT DONE
#
# Use a list comprehension with two `for` clauses (a nested loop) to build
# `pairs`: every (row, col) tuple where row is in `rows` and col is in `cols`.
#
# Nested-loop comprehension: [expr for a in A for b in B]
# Example: [(x, y) for x in [1, 2] for y in ['a', 'b']]
#   →  [(1, 'a'), (1, 'b'), (2, 'a'), (2, 'b')]

rows = [0, 1]
cols = [0, 1, 2]

pairs = ???
