# Exercise: Comprehensions 5
# I AM NOT DONE
#
# Use a dict comprehension to build `squared`: a mapping from each number in
# `numbers` to its square.
#
# A dict comprehension: {key_expr: val_expr for item in iterable}
# Example: {x: x**2 for x in [1, 2, 3]}  →  {1: 1, 2: 4, 3: 9}

numbers = [1, 2, 3, 4, 5]

squared = ???
