# Exercise: Comprehensions 6
# I AM NOT DONE
#
# Use a set comprehension to build `unique_lengths`: the set of distinct word
# lengths in `words`.
#
# A set comprehension: {expr for item in iterable}
# Example: {len(w) for w in ["hi", "hey", "yo"]}  →  {2, 3}

words = ["apple", "bat", "car", "dog", "elephant", "fig"]

unique_lengths = ???
