# Exercise: Comprehensions 9
# I AM NOT DONE
#
# Use a list comprehension to flatten `matrix` (a list of lists) into a single
# list called `flat`.
#
# Flattening with a comprehension uses two `for` clauses:
# [item for sublist in outer for item in sublist]
# Example: [[1, 2], [3, 4]]  →  [1, 2, 3, 4]

matrix = [[1, 2, 3], [4, 5, 6], [7, 8, 9]]

flat = ???
