# Exercise: Classes 5
# I AM NOT DONE
#
# Add an `increment` method to Counter that adds 1 to `self.count`.
# The method currently does nothing — replace `pass` with the right line.

class Counter:
    def __init__(self):
        self.count = 0

    def increment(self):
        pass


c = Counter()
c.increment()
c.increment()
c.increment()
