# Exercise: Tuples 5
# I AM NOT DONE
#
# Tuples are immutable — you cannot change an element in place.
# Instead, build a NEW tuple called `updated` that is the same as
# `original` but with the second element (index 1) replaced by 99.
#
# Expected: updated == (10, 99, 30)

original = (10, 20, 30)

updated = original  # fix this — do not modify original, build a new tuple
