# Exercise: Tuples 10
# I AM NOT DONE
#
# You are given a list of (name, score) tuples representing student results.
#
# Complete the code so that:
#   `top_student` — the name of the student with the highest score
#   `average`     — the average score (total / number of students), as a float
#   `passed`      — a list of names whose score is >= 60, in the original order
#
# Expected:
#   top_student == "Diana"
#   average     == 71.0
#   passed      == ["Alice", "Diana", "Eve"]

results = [
    ("Alice", 85),
    ("Bob", 42),
    ("Carol", 55),
    ("Diana", 91),
    ("Eve", 82),
]

top_student = ???
average = ???
passed = ???
