# Exercise: Tuples 9
# I AM NOT DONE
#
# Tuples can be used as dictionary keys because they are hashable.
# Complete the dictionary `grid` so that:
#   - the key (0, 0) maps to "origin"
#   - the key (1, 0) maps to "right"
#   - the key (0, 1) maps to "up"
#
# Then look up `grid[(1, 0)]` and store the result in `label`.

grid = {
    ???: "origin",
    ???: "right",
    ???: "up",
}

label = ???
