# Exercise: Tuples 8
# I AM NOT DONE
#
# Iterate over the tuple `scores` and collect every score that is
# greater than or equal to 60 into a list called `passing`.
#
# Expected: passing == [85, 92, 60, 74]

scores = (45, 85, 92, 55, 60, 74)

passing = []
for score in scores:
    pass  # append score to passing if it is >= 60
