# Exercise: Tuples 6
# I AM NOT DONE
#
# Complete the function `min_max` so that it returns a tuple of two
# values: (minimum, maximum) of the given list.
#
# Do not use multiple return statements — return both values as a
# single tuple in one `return` statement.


def min_max(numbers):
    pass  # replace with a return statement that returns (min, max)


result = min_max([4, 1, 7, 2, 9, 3])
