# Exercise: Functions 1
# ----------------------
# I AM NOT DONE
#
# Fix the function below so it takes two numbers and returns their average.

def average():
    return (a + b) / 2

# --- checks (do not edit below) ---
assert average(2, 4) == 3
assert average(10, 20) == 15
assert average(-2, -4) == -3
assert average(-10, -20) == -15
assert average(1.5, 2.5) == 2
assert average(0.5, 1.5) == 1
assert average(0, 0) == 0
assert average(3, 4.5) == 3.75
print("functions1 ✓")
