# Exercise: Functions 8
# I AM NOT DONE
#
# Define a function `build_profile` that accepts any number of keyword
# arguments using `**kwargs` and returns them as a plain dict.
# Example: build_profile(name="Alice", age=30) -> {"name": "Alice", "age": 30}
# Example: build_profile() -> {}

def build_profile(**kwargs):
    pass
