# Exercise: Generators 4
# I AM NOT DONE
#
# Use a generator expression (not a list comprehension) to build a lazy
# sequence of cubes: 1³, 2³, 3³, ..., 5³.
#
# Store the generator expression in `cube_gen`.
# A generator expression looks like: (expr for var in iterable)


# Replace None with a generator expression.
cube_gen = None
