# Exercise: Generators 1
# I AM NOT DONE
#
# Complete `count_up_to(n)` so it is a generator function that yields
# the integers 1, 2, ..., n one at a time.
#
# Hint: replace the `pass` with a `yield` statement inside the loop.


def count_up_to(n):
    for i in range(1, n + 1):
        # yield the value here
        pass
