# Exercise: Collections 8
# -----------------------
# I AM NOT DONE
#
# A deque makes an efficient FIFO queue: enqueue with `append` (right)
# and dequeue with `popleft` (left).
#
# Simulate a print queue:
#   - Enqueue jobs "doc1", "doc2", "doc3" (in that order).
#   - Dequeue them one by one into `first`, `second`, `third`.
# The deque `queue` should be empty afterwards.

from collections import deque

queue = deque()

# enqueue three jobs
???

# dequeue in FIFO order
first = ???
second = ???
third = ???
