# Exercise: Collections 5
# -----------------------
# I AM NOT DONE
#
# `namedtuple` creates a lightweight class whose instances are tuples
# with named fields — readable without index magic.
#
# Define a namedtuple type called `Point` with fields `x` and `y`.
# Then create an instance `p` with x=3 and y=4.

from collections import namedtuple

Point = ???
p = ???
