# Exercise: Collections 1
# -----------------------
# I AM NOT DONE
#
# Use collections.Counter to count how many times each word appears
# in `words`. Store the resulting Counter in `counts`.

from collections import Counter

words = ["apple", "banana", "apple", "cherry", "banana", "apple"]
counts = ???
