# Exercise: Exceptions 1
# I AM NOT DONE
#
# Complete safe_divide(a, b) so that dividing by zero is caught and
# the function returns the string "error" instead of crashing.
# A normal division should return the numeric result.
#
# Wrap the body in a try/except block:
#   try:
#       ...
#   except ZeroDivisionError:
#       ...

def safe_divide(a, b):
    return a / b
