# Exercise: Exceptions 6
# I AM NOT DONE
#
# Complete require_positive(n) so that:
# - If n is greater than 0, return n.
# - Otherwise, raise a ValueError.
#
# Use the `raise` statement to throw an exception:
#   raise ValueError

def require_positive(n):
    # TODO: raise ValueError when n is not positive
    return n
