# Exercise: Exceptions 2
# I AM NOT DONE
#
# Complete parse_number(s) so that it converts the string `s` to an int
# and returns it. If `s` cannot be converted (i.e. int(s) raises a
# ValueError), return -1 instead.
#
# Catch the *specific* exception:
#   except ValueError:

def parse_number(s):
    return int(s)
