# Exercise: Exceptions 7
# I AM NOT DONE
#
# Complete check_age(age) so that:
# - If age is between 0 and 120 (inclusive), return age.
# - Otherwise, raise a ValueError with the message:
#   "age must be between 0 and 120"
#
# Pass a message string when raising:
#   raise ValueError("your message here")

def check_age(age):
    # TODO: raise ValueError with the message "age must be between 0 and 120"
    # when age is out of range
    return age
