# Exercise: Testing 4
# I AM NOT DONE
#
# A test function can (and usually should) contain *multiple* assertions —
# one for each behaviour you want to verify.
#
# Complete `test_is_even` so it asserts ALL of the following:
#   is_even(0)  is True
#   is_even(2)  is True
#   is_even(7)  is False
#   is_even(-4) is True

def is_even(n):
    return n % 2 == 0


def test_is_even():
    # write four assertions here
    pass
