# Exercise: Testing 3
# I AM NOT DONE
#
# Tests are usually wrapped in a function whose name starts with `test_`.
# When the function is called it either runs to completion (pass) or raises
# an AssertionError (fail).
#
# Complete `test_add` so it asserts:
#   add(2, 3) == 5
#   add(0, 0) == 0
# A correct test runs without raising.

def add(a, b):
    return a + b


def test_add():
    # write assertions here
    pass
