# Exercise: Testing 1
# I AM NOT DONE
#
# The simplest test in Python is a plain `assert` statement.
# If the expression is True, nothing happens; if False, an AssertionError
# is raised and the test fails.
#
# Replace ??? so that the assertion holds:
#   result should equal 4 after calling double(2).

def double(x):
    return x * 2


result = double(2)
assert result == ???
