# Exercise: Testing 2
# I AM NOT DONE
#
# An assert can carry an optional message that is shown when it fails:
#   assert condition, "helpful message"
#
# Replace ??? with the correct expected value AND complete the message
# so the assertion has the form:
#   assert square(5) == ???, "5 squared should be 25"

def square(x):
    return x * x


assert square(5) == ???, "???"
